# Assignment No 2c
#### Assignment Credit 
###### *Dr. Sibt ul Hussain*


Text Classification (Sentiment Analysis) Using Bayes Rule 
==============


## Goal

Your goal in this part of assigment is to implement a Naive Bayes Multinomial classifier using  bag of words model for the classification of text (movie reviews) into different categories..

**Note** Please note that you are allowed to use only those libraries which we have discussed in the class, i.e. numpy, scipy, pandas.

Once you have build and test the model on the provided dataset. You will use the learned techniques to compete in a [Kaggle](https://www.kaggle.com/c/word2vec-nlp-tutorial) competition and report your final score and leaderboard ranking to get full credit.

For final submission attach the screen-shot of the leader-board with your score.

## Submission Instructions
You are required to submit the original notebook file on the Slate (with .ipynb extension), with complete set of outputs. Students failing to do so will get zero marks. 

*Please read each step carefully and understand it fully before proceeding with code writing*

## Plagiarism
Any form of plagiarism will not be tolerated and result in 0 marks.


```python
%pylab inline
import scipy.stats
from collections import defaultdict
```

Output:
```
%pylab is deprecated, use %matplotlib inline and import the required libraries.
Populating the interactive namespace from numpy and matplotlib
```

<!-- grade_id: parse_string, points: 0, grade: false, solution: true -->
```python
import re
import bs4

def parse_string(string): 
    """"
        Parse the input string and tokenize it using regular expressisons:
        First clean the string such that it does not have any punctuation or number, it must only have a-z and A-Z.
        Please note that while doing this, the spaces much not get disturbed, but in case of multiple spaces convert 
        them to one space.
        Then convert the string to lower case and return its words as a list of strings.
        Example:
        --------
        Input :  computer scien_tist-s are,,,  the  rock__stars of tomorrow_ <cool>  ????
        Output:  ['computer', 'scientists', 'are', 'the', 'rockstars', 'of', 'tomorrow']
        
        Parameters:
        ----------
        string: string to be parsed...
        re: regular expression to be used for the tokenization.
        
        Returns:
        ---------
        list of tokens extracted from the string...
    """
  
    
    # YOUR CODE HERE
    #can we use string library or nltk?
    #raise NotImplementedError()
    #replace multiple spaces with just one
    
    string = re.sub(r'<[^>]+>', '', string)
    # Removing punctuations in string
    # Using regex
    string = re.sub(r'[^\w\s]|_', '', string)
    string = re.sub(r'\s+', ' ', string)
    #string = re.sub(r'[!"#$%&\'()*+,-./:;<=>?@[\\]^_`{|}~]','', string)
    l = string.lower().split()
    #print(l)
    return l
#parse_string("co)*+,-./:;<=>?@[\\]mputer  $~  scien_ti%&\'(^_`{|}st-s are,,,                    the  rock__stars of tomorrow_ <cool>  ????")
```

<!-- grade_id: parse_file, points: 0, grade: false, solution: true -->
```python
def parse_file(filename): # Parse a given file
    """"
        Parameters:
        ----------
        filename: name of text file to be read
   
        
        Returns:
        ---------
        read file as raw string (with \n, \t, \r, etc included)
    """
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    with open(filename, 'r') as f:
        
        content = f.read() #parse_string(f.read()) #if not preprocess here then when?
        
    
    return content

#f = parse_file(r"C:\Users\mahee\Downloads\Fall24_assignment-2\Fall24_assignment-2\data\imdb1\pos\cv997_5046.txt")
```

<!-- grade_id: files_to_string, points: 0, grade: false, solution: true -->
```python
def files_to_strings(X): #converts file to string
    
    """
        Read an array (or list) of files where each file content is read in a string...
        Input:
        -------
        X an array (or list) of file names
        
        Returns:
        --------
        X as a numpy array with each row containing a read string from the file...
    """
    
    # YOUR CODE HERE
    #raise NotImplementedError()
    
    l = []
    
    for file in X:
        l.append(parse_file(file))
    l = np.array(l, dtype = object) 
    #print(l.shape)
    return l

# strs = files_to_strings(np.array(["./data/imdb1/neg/cv000_29416.txt", "./data/imdb1/pos/cv000_29590.txt"]))
# print(strs[1])

# stri = parse_string(strs[1])
# print(stri)

```

<!-- grade_id: test_reading_and_parsing, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_list_equal

assert_list_equal(parse_string("computer scien_tist-s are,,,  the  rock__stars of tomorrow_ <cool>  ????"),
        [u'computer', u'scientists', u'are', u'the', u'rockstars', u'of', u'tomorrow'], "Incorrect cleanning")


strings = files_to_strings(np.array(["./data/imdb1/neg/cv000_29416.txt", "./data/imdb1/pos/cv000_29590.txt"]))
with open("./data/imdb1/neg/cv000_29416.txt") as f:
    text = f.read()
assert_equal(strings[0], text, "At first index should be text of first file") # doesn't make sense for this to be here,
#strings[0] will be cleaned(one you processed while text has just been read)
assert_equal(strings.shape, (2,), "Shape must be (2,) for two files in list")
```

Output:
```
File c:\ProgramData\anaconda3\Lib\site-packages\nose\loader.py:21
     19 from nose.failure import Failure
     20 from nose.config import Config
---> 21 from nose.importer import Importer, add_path, remove_path
     22 from nose.selector import defaultSelector, TestAddress
     23 from nose.util import func_lineno, getpackage, isclass, isgenerator, \
     24     ispackage, regex_last_key, resolve_name, transplant_func, \
     25     transplant_class, test_address

File c:\ProgramData\anaconda3\Lib\site-packages\nose\importer.py:12
      9 import sys
     10 from nose.config import Config
---> 12 from imp import find_module, load_module, acquire_lock, release_lock
     14 log = logging.getLogger(__name__)
     16 try:

ModuleNotFoundError: No module named 'imp'
```

<!-- grade_id: classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
        
## Hint, you can use python dictionary or default dict for counting the words
# or counter class from collections 

#TODO Complete this class for running the complete classifier... 

#You might need to define auxiliary classes for the complete algorithm..

class NaiveBayes:
    ''' Implements the Naive Bayes For Text Classification... '''
    def __init__(self, classes):
        self.classes=classes
        self.p_prior = dict.fromkeys(self.classes, 0)
        self.uniq = []

        
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        
    def addExample(self, x, y):
        '''
            Add example to corresponding class model ...
            Input
            ---------
            x: example (list of words)
            y: label...
        '''
        
        # YOUR CODE HERE
        #??? what example?
        raise NotImplementedError()
        
    def train(self, X, Y):
        ''' Train the multiclass (or Binary) Bayes Rule using the given 
            X [m x d] data matrix and Y labels matrix
            
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional examples.
            Y: [m x 1] a label vector.
            
            Returns:
            -----------
            Nothing
            '''
        #how will it learn???
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        l = []

        for i in X: #for each file in X get cleaned string list
            l.append(parse_string(str(i)))
        X = np.array(l, dtype = object)
        
        #create bag of words (tokenize words and count frequency of each word) #bag of words for each doc or all docs?
        
        j = ''
        u = []
        for i in X:  # file i in X
            for j in i: #word j in file i
                u.append(j)
        self.uniq = np.array(list(set(u))) #list of all features
        bow = dict.fromkeys(self.uniq, 0) # dictionary for bow(one bag for all docs)
        for i in X:
             #count frequency
            for j in i:
                bow[j] += 1
            
        
        #already in unigrams as unigrams contain just words

        #labels = list(set(Y))
        
        
        for i in Y:       #to calculate prior probability 
            self.p_prior[i] += 1
        for key, val in self.p_prior.items():  
                                            # how to calculate word/label? number of times word appears in all docs/total words in all docs?
            self.p_prior[key] /= len(Y)
        print(self.p_prior)
            
        
        #create uni-gramsas we will be using naive bayes
        #probability of class given features(words)
        
        #probability(C/doc) = P(word/C) * P(word2/doc)....
        #use MAP (Maximum A Posteriori) p(c/X) = P(X1/C) * P(X2/C) * ...*P(C)  
        #c = argmax(probs)   
        
    def test(self, X):
        
        ''' Test the trained classifiers on the given set of examples 
        
                   
            Input:
            ------
            X: [m x d] a data matrix of m d-dimensional test examples.
           
            Returns:
            -----------
                pclass: the predicted class for each example, i.e. to which it belongs
        '''
        
        nexamples, nfeatures=X.shape

        l = [] 

        for i in X: #for each file in X get cleaned string list
            l.append(parse_string(str(i)))
        X = np.array(l, dtype = object) # each file is in a list of words

        # each class has it's own dic and this dic contains all features with their counts 
        p_evid = dict.fromkeys(self.classes, {}) #calculate likelihood for c1 and c2
    
        j = ''
        u = []
        probs = []
        for i in X:  # file i in X
            for j in i: #word j in file i
                u.append(j)
            unique = np.array(list(set(u))) #list of all features
            for l in self.classes: # against every label add words are labels with 0 as value
                p_evid[l] = dict.fromkeys(unique, 0)
            
            j = ''
            for file in X: 
                for word in file:
                    for c in self.classes:
                        p_evid[c][word] += 1 #now you have counts of features belonging to class
            
            pcw = dict.fromkeys(self.classes, 1) #probability of class given word #values set 1 so that they don't become 0
            
            for label, words in p_evid.items(): #label, dict from p_evid
                for w,c in words.items():
                    
                    #all words belonging to that class
                    p_evid[label][w] /= sum(p_evid[label][w]) #likelihood probabilities calculated words
                                        #total words
                    #get index of label
            #print(p_evid)        
            for label, words in p_evid.items():
                for w,c in words.items(): #no likelihood is 0 here
                    pcw[label] *= p_evid[label][w] #multiply likelihoods
                pcw[label] *= self.p_prior[label] #multiply prior probability   
            
            for l, p in pcw.items():
                print('label: ', l, "probability: ", p)
            max_key = max(pcw, key=pcw.get) #get key with maximum posterior probability
            prob.append(max_key)
          
        # YOUR CODE HERE
        #raise NotImplementedError()
        return
        
    def predict(self, x):
        '''
            Predict the label of given input example...
            
            Input
            ---------
            x: example (list of words)
            
        '''


        
        # YOUR CODE HERE
        #raise NotImplementedError()
        
        return
    
X=files_to_strings(tX) # read files and convert each file into set of strings and return an numpy array
X = X.reshape((X.shape[0], 1))
#Split the data into two halves training and test set...
traindata,trainlabels,testdata,testlabels=t.split_data(X,labels)
#Find the classes to train
classes=np.unique(labels)
nb = NaiveBayes(classes)
nb.train(traindata,trainlabels)
nb.test(testdata)
```

Output:
```
{'neg': 0.5, 'pos': 0.5}
label:  neg probability:  0.5
label:  pos probability:  0.5
```

```python
import pandas as pd
import tools as t
```

```python
tdir= "./data/imdb1/" # training dir...
#load data, get list of files for each class...
posfiles=t.get_files(tdir+'/pos','*',withpath=True)
negfiles=t.get_files(tdir+'/neg','*',withpath=True)
```

Output:
```
1000
```

```python
#generate training and testing data...
plabels=['pos']*len(posfiles)
nlabels=['neg']*len(negfiles)
labels=np.concatenate((plabels,nlabels)) # concatenate the +ve and -ve labels
tX=np.concatenate((posfiles,negfiles))
print ("Training data Dimensions =", tX.shape," Training labels dimensions=", labels.shape)
```

Output:
```
Training data Dimensions = (2000,)  Training labels dimensions= (2000,)
['./data/imdb1//neg\\cv001_19502.txt' './data/imdb1//neg\\cv002_17424.txt']
```

```python
X=files_to_strings(tX) # read files and convert each file into set of strings and return an numpy array
X = X.reshape((X.shape[0], 1))
#Split the data into two halves training and test set...
traindata,trainlabels,testdata,testlabels=t.split_data(X,labels)
#Find the classes to train
classes=np.unique(labels)
```

Output:
```
[['films adapted from comic books have had plenty of success , whether they\'re about superheroes ( batman , superman , spawn ) , or geared toward kids ( casper ) or the arthouse crowd ( ghost world ) , but there\'s never really been a comic book like from hell before . \nfor starters , it was created by alan moore ( and eddie campbell ) , who brought the medium to a whole new level in the mid \'80s with a 12-part series called the watchmen . \nto say moore and campbell thoroughly researched the subject of jack the ripper would be like saying michael jackson is starting to look a little odd . \nthe book ( or " graphic novel , " if you will ) is over 500 pages long and includes nearly 30 more that consist of nothing but footnotes . \nin other words , don\'t dismiss this film because of its source . \nif you can get past the whole comic book thing , you might find another stumbling block in from hell\'s directors , albert and allen hughes . \ngetting the hughes brothers to direct this se
```

```python
#Now build a Naive Bayes classifier and test it...
print ('[Info] training a classifier for following classes {}, {}'.format(classes[0],classes[1]))
nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
pclasses=nb.test(testdata)
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
print ("[Info] Accuracy = {}".format(acc))
```

Output:
```
[Info] training a classifier for following classes neg, pos
{'neg': 0.5, 'pos': 0.5}
Cell In[33], line 5
      3 nb=NaiveBayes(classes)
      4 nb.train(traindata,trainlabels)
----> 5 pclasses=nb.test(testdata)
      6 acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
      7 print ("[Info] Accuracy = {}".format(acc))

Cell In[27], line 142, in NaiveBayes.test(self, X)
    139 nexamples, nfeatures=X.shape
    141 # YOUR CODE HERE
--> 142 raise NotImplementedError()
    143 return

NotImplementedError:
```

### Test Cells Start
#### Do not Modify

<!-- grade_id: test_classifier_shapes, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_in

nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
assert_equal (nb.test(testdata).shape[0], testdata.shape[0])
assert_in( type(nb.predict(["ok"])) , [str, np.string_, np.str, np.str_] , "Predict should return a label \
                                                                                            not list or array")
```

<!-- grade_id: test_acc, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_greater

nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
pclasses=nb.test(testdata)
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
assert_greater(acc, 0.77, "Acc must be greater then 77% you are doing something wrong")    
```

<!-- grade_id: test_classifier_responce, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_equal

comment_pos = "A nice movie, the case was good. Overall a perfect play"
comment_neg = "A waste of time, cast was bad. a clear No!"

#generate training and testing data...
tX=np.concatenate((posfiles,negfiles))
X=files_to_strings(tX)
X = X.reshape((X.shape[0], 1))

plabels=['pos']*len(posfiles)
nlabels=['neg']*len(posfiles)
true_labels = np.concatenate((plabels,nlabels))
inverted_labels = np.concatenate((nlabels,plabels))

true_nb=NaiveBayes(classes)
true_nb.train(X,true_labels)

inverted_nb=NaiveBayes(classes)
inverted_nb.train(X,inverted_labels)

assert_equal( true_nb.predict(comment_pos.split()), "pos" )
assert_equal( true_nb.predict(comment_neg.split()), "neg" )

assert_equal( inverted_nb.predict(comment_pos.split()), "neg" )
assert_equal( inverted_nb.predict(comment_neg.split()), "pos" )
```

### Test Cells End

# Cross Validation

Now lets throw our methods to winds of different folds and measure their accuracy...

```python
#Now lets generate n-fold training and testing data...
nfolds=10
folds=t.generate_folds(X,labels,nfolds) # generate folds for 
for k in arange(len(folds)):
    print (folds[k][0].shape, folds[k][2].shape)
```

```python
totacc=[]
#train a classifier for each fold...
classes=np.unique(labels)

for k in range(nfolds):
    nb=NaiveBayes(classes)
    
    traindata=folds[k][0]
    trainlabels=folds[k][1]
    
    #Lets first train the classifier
    nb.train(traindata,trainlabels)
    
    testdata=folds[k][2]
    testlabels=folds[k][3]
    
    #Lets test the classifier
    pclasses= nb.test(testdata)
    
    #print pclasses
    acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
    print ("[Info] Fold {} Accuracy = {}".format(k+1, acc))
    
    totacc.append(acc)

print (totacc)

mean_acc = np.mean(totacc)
print ('[Info] Mean Accuracy =', mean_acc)
```

# Excellent, now its time to go into real waters of Kaggle.


You will be needed to create an account on the Kaggle and download the data for the competition ["Bag of words meets bags of popcorn"](https://www.kaggle.com/c/word2vec-nlp-tutorial/data).  Note that you will be only downloading the "labeledTrainData.tsv" and "labeledTestData.tsv".


"labeledTrainData.tsv" will be used for training your model and thus have prespecified labels for each example review. "labeledTestData.tsv" will be used for testing your model and thus don't have prespecified labels for each example. You will predicting the label for each review and then uploading your result to Kaggle server which will be evaluating your model and will give score to your entry. You will report this score during your assignment submission.

**[Caution]** Please note that Kaggle limits maximum number of evaluations per 24 hours to 5 to reduce the overfitting on the test set, so be careful and throughly test your model before submitting your entry to Kaggle server. 

Read the instructions on the Competition Page. Note you are not allowed to use any of the library except what we have learned during class.

```python
# read the data-set
train=pd.read_csv('/tmp/labeledTrainData.tsv',sep='\t')
```

```python
train.describe()
```

```python
train.head()
```

```python
Yt=train['sentiment']
Xt=train['review']
Xt=np.array(Xt)
Yt=np.array(Yt)

print (Xt.shape)
```

```python
#read test set...
test=pd.read_csv('/tmp/testData.tsv',sep='\t')
```

```python
test.head()
```

#### Training Time 

```python
# Let's split the training data into two halves and test our accuracy...
traindata,trainlabels,testdata,testlabels=t.split_data(Xt.reshape((Xt.shape[0],1)),Yt)
classes=np.unique(trainlabels)
```

```python
# Now lets go and train the model and see its performance...
print ('[Info] training a classifier for following classes {}, {}'.format(classes[0],classes[1]))
nb=NaiveBayes(classes)
nb.train(traindata,trainlabels)
pclasses=nb.test(testdata)
acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
print ("[Info] Accuracy = {}".format(acc) )
```

#### Cross-Validation Time...

```python
#Split the training data into 10 folds and test classifiers performance...

nfolds=10
folds=t.generate_folds(Xt.reshape((Xt.shape[0],1)),Yt,nfolds) # generate folds for 
for k in arange(len(folds)):
    print (folds[k][0].shape, folds[k][2].shape)
```

```python
# As it takes time, so becareful it can cause your machine into red hot oven
totacc=[]
classes=np.unique(Yt)

for k in range(nfolds):
    nb=NaiveBayes(classes)
    
    traindata=folds[k][0]
    trainlabels=folds[k][1]
    
    #Lets first train the classifier
    nb.train(traindata,trainlabels)
    
    testdata=folds[k][2]
    testlabels=folds[k][3]
    
    #Lets test the classifier
    pclasses= nb.test(testdata)
    
    acc=np.sum(pclasses==testlabels)/float(testlabels.shape[0])
    print ("[Info] Fold {} Accuracy = {}".format(k+1, acc) ) 
    
    totacc.append(acc)

print (totacc)
print ('[Info] Mean Accuracy =', np.mean(totacc))
```

# Now let's train on the complete dataset and test on the provided test set...

```python
classes= np.unique(Yt)
print ('Training a Classifier on Full training set with classes =', classes)
nb=NaiveBayes(classes)
nb.train(Xt.reshape(Xt.shape[0],1),Yt)
```

```python
#Get the test data...
Xtest=test['review']
Xtest=np.array(Xtest.reshape((Xtest.shape[0],1)))
#test the classifier on the provided test set...
pclasses=nb.test(Xtest)
```

```python
#write the result in the kaggle's required format
output = pd.DataFrame( data={"id":test["id"], "sentiment":pclasses} )

# Use pandas to write the comma-separated output file
output.to_csv( "/tmp/Naive_bays_Bag_of_Words_model.csv", index=False, quoting=3 )
```

# Time to Upload the prediction to Kaggle...

Now upload the result on the Kaggle and see your ranking and score. Using this simple method you can have an accuracy of around 0.80960.

# Improvement by Excluding Stop Words...

You can improve your score further by excluding the commonly occuring words (also known as stop words) in the English language.



```python
#read and create a set of stop 
stopwords=set(t.read_txt_file('./data/sentiment/data/english.stop'))
print (stopwords)
```

Now you can re-build the model by excluding these words and again upload your results on Kaggle. 

Doing this simple trick can further improve your accuracy to 0.81768.

For final submission attach the screen-shot of the leader-board with your score

Insert ScreenShot of Leader-board Below
----------------------------------------
